# <span style="color:teal;">CIS 211 Project 8:  &nbsp; Bears and Fish</span>

##### Due 11:00 P.M. Wednesday Mar 9

##### <span style="color:red">Group Members:</span>

Group work is allowed for this project.  Write the name and DuckID for each member of the group in this markdown cell:

Name: Amber Straub 951333723

Groups are exepcted to do more work.  See the sections at the end of this notebook for suggestions of how to extend the simulation.

##  <span style="color:teal;">Overview</span> 

The last project this term is a cellular automaton style biological simulation.  

The sytem is described in Chapter 11 of the Miller and Ranum text.  The simulation needs three class defintions:
* World defines a 2D grid where cells are accessed according to their row and column coordinates
* Fish defines one type of animal that can inhabit the world; fish can move and reproduce, and they can die from overcrowding
* Bear defines the other type of animal; bears can also move and reproduce, they eat fish, and they can die from starvation

**<span style="color:red">Hint:</span>** &nbsp; Before you start, think about what the Bear and Fish classes have in common, and consider defining a new class to use as the parent of both Bear and Fish.

##  <span style="color:teal;">Libraries</span> 

Your program will need to use `numpy` (a numeric processing library) and `random` (the builtin random number generator module).  You can import additional modules if you wish; add the import statements to this code cell.

```python
import numpy as np
import random
```

##  <span style="color:teal;">Event Log</span> 

A useful debugging technique is to save descriptions of events in a log file.  For this project we'll just print messages so they appear in the output area of the code cell that calls the top level simulator.

If you want to see messages printed as your program runs set `logging` to True in the code cell below, and include calls to `log` at the appropriate places in your methods.

**Note:** &nbsp; you can change the definition of the `log` function so it writes to a file or  prints additional information.

```python
logging = False

def log(message):
    if logging:
        print(message)
```

##  <span style="color:teal;">Classes, Functions, and Global Variables</span> 

As you write your code you can define additional classes, helper functions, and global variables.  Describe all of new classes, functions and variables in the following markdown cell, and put all of the definitions in one code cell right below the documentation.

**************
To simplify the code for my Bear and Fish class, I have created an extra class that includes the traits they have in common. I explain what it's functions are in my Fish and bear classes.

***********

```python
class Animal():
    def __init__(self, world, tuple):
        row_pos,column_pos=tuple
        self.world=world
        self.breedTick=0
        self.column_pos=column_pos
        self.row_pos=row_pos
        self._alive=True
        
        if self not in world.biota():
            world.addThing(self)
   

    def breedTick(self):
        return self.breedTick
    def set_row(self, new_row):
        self.row_pos=new_row
    def column_pos(self, new_column):
        self.column_pos=new_column
    def get_row(self):
        return self.row_pos
    def get_column(self):
        return self.column_pos
    
       
    def do_move(self, new_row, new_column):
        self.world.moveThing(self.row_pos, self.column_pos, new_row, new_column)
        self.row_pos=new_row
        self.column_pos=new_column
       
    
   
    def move(self):
        offA=(-1,0,1,-1,1,-1,0,1)
        offB=(1,1,1,0,0,-1,-1,-1)
        randomOffsetIndex=random.randrange(len(offA))
        randA=offA[randomOffsetIndex]
        randB=offB[randomOffsetIndex]
        new_row=self.row_pos+randA
        new_column=self.column_pos+randB
        while not (new_row>=0, new_column>=0, new_row<=self.world.row, new_column<=self.world.column):
            randomOffsetIndex=random.randrange(len(offA))
            randA=offA[randomOffsetIndex]
            randB=offB[randomOffsetIndex]
            new_row=self.row_pos+randA
            new_column=self.column_pos+randB
        if self.world.knock(new_row, new_column)==None:
            self.do_move(new_row, new_column)
            
    def location(self):
        return (self.row_pos, self.column_pos)
    
    def tryToBreed(self, thing):
        
        offA=(-1,0,1,-1,1,-1,0,1)
        offB=(1,1,1,0,0,-1,-1,-1)
        randomOffsetIndex=random.randrange(len(offA))
        randA=offA[randomOffsetIndex]
        randB=offB[randomOffsetIndex]
        new_row=self.row_pos+randA
        new_column=self.column_pos+randB
        while not (new_row>=0, new_column>=0, new_row<=self.world.row, new_column<=self.world.column):
            randomOffsetIndex=random.randrange(len(offA))
            randA=offA[randomOffsetIndex]
            randB=offB[randomOffsetIndex]
            new_row=self.row_pos+randA
            new_column=self.column_pos+randB
        if self.world.knock(new_row, new_column)==None:
            child=thing(self.world, (new_row, new_column))
            if thing==Bear:
                self.breedTick=1
            else:   
                self.breedTick=0 
```

##  <span style="color:teal;">World</span>

Describe your World class in the following markdown cell, and write the code for the class in the code cell below the markdown cell.

The minimum requiremens for this class are:
* the constructor should be passed the grid size (number of rows and columns)
* all cells should contain `None` when the world is initialized; later they can contain references to Fish or Bear objects
* overload the methods that allow you to access a cell or store a value in a cell
* include a method named `biota` that returns a list of Fish and Bear objects currently in the grid

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

**Note:** &nbsp; There are no auto-grader tests for the World class; tests for this class will be done as part of the tests for the Bear and Fish classes.

<!-- grade_id: world_doc, points: 10, grade: true, solution: true -->
*****************
   Below is my implementation of the world class. This class is a numpy two dimentional array with columns and rows. Each tile in the array can be accessed and edited. Items can be added to the array, and once added are accessable and mutable. We can move, add and delete object from this array.
***************


<!-- grade_id: world, points: 15, grade: true, solution: true -->
```python
class World():
    def __init__(self, row, column):
        self.row=row
        self.column=column
        self.itemList=[]
        self.grid=np.array([None]*(row*column)).reshape(row, column)
        
    def row(self):
        return self.row
    def column(self):
        return self.column
    
    
    def biota(self):
        return self.itemList
    
    def grid(self):
        return self.grid
        
    def knock(self, row, column):
        return self.grid[column][row]
        
    
    def addThing(self, thing):
        
        self.grid[thing.row_pos][thing.column_pos]=thing
        self.itemList.append((thing))
        
        
    def delThing(self, thing):
        
        self.grid[thing.row_pos][thing.column_pos]=None
        self.itemList.remove((thing))
        
    def moveThing(self, oldrow, oldcolumn, newrow, newcolumn):
        self.grid[newrow][newcolumn]=self.grid[oldrow][oldcolumn]
        self.grid[oldrow][oldcolumn]=None
    
    def live(self):
        if self.itemList!=[]:
            thing=random.randrange(len(self.itemList))
            randomthing=self.itemList[thing]
            randomthing.live()
  
        
```

### <span style="color:teal;">Fish</span>

Use the markdown cell and code cell below to implement and describe your Fish class.

During the simulation a Fish object needs to behave as follows:

(1) Fish are susceptible to overcrowding:  if there are fish in 2 or more neighboring cells the fish dies (it's removed from the simulation)

(2) A fish can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new fish is placed in that cell

(3) A fish can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the fish moves there

The constructor will be passed a reference to a World object and a location, in the form of a tuple with a row number and column number (the object needs to know its location so it can look for other objects in neighboring cells).

The class should include the following methods:
* `live` implements rules 1 and 2 shown above
* `move` implements rule 3
* `location` returns the current grid location (row and column) of the object

Define a class variable named `breed_interval` to specify how many time steps a fish must be alive before it reproduces; the initial value for this variable is 12.

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

**NOTE:** &nbsp; The way the simulation is defined a fish might be eaten before the top level simulation calls the `live` method.  Your method should check to make sure the fish is still alive.  The easiest way to do this is to include an instance variable named `_alive` that is set to True when the fish is initialized and set to False when it dies.

<!-- grade_id: fish_doc, points: 5, grade: true, solution: true -->
******
This is my implementation of the Fish class. It inherits most of it's attribites from it's parent Animal class. These attributes include move, it's location, and it's breeding specifications. The most pertinent attribute below is the live function, which tells the fish whether it starves, dies from old age, eats a plant, or gets to attempt breeding. Fun stuff.
*******

<!-- grade_id: fish, points: 10, grade: true, solution: true -->
```python
class Fish(Animal):
    
    breed_interval=12
    survive_without_food=5
    def __init__(self, world, tuple):
        Animal.__init__(self, world, tuple)
        self.surviveTic=1
   
    def __repr__(self):
        return "Fish"

    def live(self):
        #Did the fish starve?
        if self.surviveTic>=Fish.survive_without_food:
            self.world.delThing(self)
        
        offA=(-1,0,1,-1,1,-1,0,1)
        offB=(1,1,1,0,0,-1,-1,-1)
        
         #is it overcrowded
        offA=(-1,0,1,-1,1,-1,0,1)
        offB=(1,1,1,0,0,-1,-1,-1)
        adjfish=0
        for offset in range(len(offA)):
            new_row=self.row_pos+offA[offset]-1
            new_column=self.column_pos+offB[offset]-1
            if new_row>=0 and new_row<=self.world.row and new_column>=0 and new_column<=self.world.column:
                if (self.world.knock(new_row, new_column)!=None and isinstance(self.world.knock(new_row, new_column), Fish)):
                    adjfish+=1
        if adjfish>=2:
            self.world.delThing(self)
        
        
        #Find food
        adjplant=[]
        for offset in range(len(offA)):
            new_row=self.row_pos+offA[offset]-1
            new_column=self.column_pos+offB[offset]-1
            if new_row>=0 and new_row<=self.world.row and new_column>=0 and new_column<=self.world.column:
                if (self.world.knock(new_row, new_column)!=None and isinstance(self.world.knock(new_row, new_column), Plant)):
                    adjplant.append(self.world.knock(new_row, new_column))
        
        if len(adjplant)>=2:
            #select one at rondom to eat
            x=random.randrange(len(adjplant))
            self.world.delThing(adjplant[x])
        
        elif len(adjplant)==1:
            #eat it
            
            self.world.delThing(adjplant[0])
        else:
            #went hungery
            self.surviveTic+=1
            
       
        #if not dead, breed
        self.breedTick+=1
        if (self.breedTick>=Fish.breed_interval):
            self.tryToBreed()
        
    
    def tryToBreed(self):
        Animal.tryToBreed(self,Fish)
    

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: fish_test_1, points: 3, grade: true, solution: false -->
```python
# A new world has no objects
w1 = World(5,5)
assert len(w1.biota()) == 0

# After adding a fish there should be one object
f1 = Fish(w1, (2,2))
assert len(w1.biota()) == 1

# Test the location method
assert f1.location() == (2,2)
```

<!-- grade_id: fish_test_2, points: 3, grade: true, solution: false -->
```python
# Setting breed_interval to 1 should cause a fish to reproduce when live is called
w2 = World(5,5)
Fish.breed_interval = 1
f2 = Fish(w2, (2,2))
f2.live()
assert len(w2.biota()) == 2

# Reset the interval to original value for remaining tests
Fish.breed_interval = 12
```

<!-- grade_id: fish_test_3, points: 3, grade: true, solution: false -->
```python
# Make three fish, the one in the middle should die from overcrowding
w3 = World(5,5)
f3 = Fish(w3, (2,2))
Fish(w3, (1,1))
Fish(w3, (3,3))
f3.live()
assert len(w2.biota()) == 2
```

<!-- grade_id: fish_test_4, points: 3, grade: true, solution: false -->
```python
# When a fish moves it should be within one cell of its original location
w4 = World(5,5)
f4 = Fish(w4, (2,2))
f4.move()
r, c = f4.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
```

###  <span style="color:teal;">Bear</span>

Use the markdown cell and code cell below to implement and describe your Bear class.

During the simulation a Bear object needs to behave as follows:

(1) The bear looks for fish in each adjacent cell; if it finds one or more fish it eats one at random 

(2) If a bear has not eaten for certain number of time steps it dies (it's removed from the simulation)

(3) A bear can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new bear is placed in that cell

(4) A bear can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the bear moves there

Define a class variable named `breed_interval` to specify how many time steps a bear must be alive before it reproduces; the initial value for this variable is 8.  Define another class variable named `survive_without_food` to be the number of time steps a bear can live before it dies from starvation; the initial value for this variable is 10.

The constructor will be passed a reference to a World object and a location, in the form of a tuple with a row number and column number (the object needs to know its location so it can look for other objects in neighboring cells).

The class should include the following methods:
* `live` implements rules 1, 2, and 3 shown above
* `move` implements rule 4
* `location` returns the current grid location (row and column) of the object

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

<!-- grade_id: bear_doc, points: 5, grade: true, solution: true -->
*****
Like the Fish class, the Bear class inherites most of it's attributes from the Animal class. The bear classes live function has some extra special functions including ckecking if it dies from old age, dies from starvation, or if it gets a chance to eat something. It will also attempt breeding after it's eaten(if it survives).
********

<!-- grade_id: bear, points: 10, grade: true, solution: true -->
```python
class Bear(Animal):
    
    breed_interval=8
    survive_without_food=10
    
    def __init__(self, world, tuple):
        Animal.__init__(self, world, tuple)
        self.surviveTic=1
   
    def __repr__(self):
        return "Bear"
    
    
       
    def live(self):
        #Is the Bear Alive?
        if self.surviveTic>=Bear.survive_without_food:
            self.world.delThing(self)
        
        offA=(-1,0,1,-1,1,-1,0,1)
        offB=(1,1,1,0,0,-1,-1,-1)
        
        #Find food
        adjfish=[]
        for offset in range(len(offA)):
            new_row=self.row_pos+offA[offset]-1
            new_column=self.column_pos+offB[offset]-1
            if new_row>=0 and new_row<=self.world.row and new_column>=0 and new_column<=self.world.column:
                if (self.world.knock(new_row, new_column)!=None and isinstance(self.world.knock(new_row, new_column), Fish)):
                    adjfish.append(self.world.knock(new_row, new_column))
        
        if len(adjfish)>=2:
            #select one at random to eat
            x=random.randrange(len(adjfish))
            self.world.delThing(adjfish[x])
        
        elif len(adjfish)==1:
            self.world.delThing(adjfish[0])
        else:
            self.surviveTic+=1
            
        #Bear breeds
        self.breedTick+=1
        if (self.breedTick>=Bear.breed_interval):
            self.tryToBreed()
        
    
    def tryToBreed(self):
        Animal.tryToBreed(self, Bear)
    
    
```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: bear_test_1, points: 3, grade: true, solution: false -->
```python
# Test the Bear constructor and location method
w1 = World(5,5)
b1 = Bear(w1, (1,1))
assert len(w1.biota()) == 1
assert b1.location() == (1,1)
```

<!-- grade_id: bear_test_2, points: 3, grade: true, solution: false -->
```python
# Repeat the reproduction test for Bears
w2 = World(5,5)
Bear.breed_interval = 1
b2 = Bear(w2, (2,2))
b2.live()
assert len(w2.biota()) == 2
Bear.breed_interval = 8
```

<!-- grade_id: bear_test_3, points: 3, grade: true, solution: false -->
```python
# Make fish for the bear to eat, count the number of objects after eating
w3 = World(5,5)
b3 = Bear(w3, (2,2))
Fish(w3, (1,1))
Fish(w3, (3,3))
b3.live()
assert len(w2.biota()) == 2
```

<!-- grade_id: bear_test_4, points: 3, grade: true, solution: false -->
```python
# Setting the survival limit to 1 should cause a Bear to starve 
w4 = World(5,5)
Bear.survive_without_food = 1
b4 = Bear(w4, (2,2))
b4.live()
assert len(w4.biota()) == 0
```

<!-- grade_id: bear_test_5, points: 3, grade: true, solution: false -->
```python
# Repeat the move test for bears
w5 = World(5,5)
b5 = Bear(w5, (2,2))
b5.move()
r, c = b5.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
```

##  <span style="color:teal;">Top Level Functions</span>

Fill in the body of the `wbf` function so it returns a new World object with the specified number of rows and columns and with the specified number of Bear and Fish objects at random locations.

When we grade your project we will call `wbf` to make a World object and then use the main loop (implemented by `step_system`) to run the simulation.

<!-- grade_id: wbf, points: 5, grade: true, solution: true -->
```python
def wbf(nrows, ncols, nbears, nfish):
    '''
    Return a new World object with the specified number of Bear and Fish objects.
    '''
    world=World(nrows,ncols)
    
    for i in range(nfish):
        
        x=random.randrange(nrows-1)
        y=random.randrange(ncols-1)
        
        while world.knock(x,y)!=None:
            x=random.randrange(nrows-1)
            y=random.randrange(ncols-1)
        child=Fish(world, (x, y))
    for i in range(nbears):
        
        x=random.randrange(nrows-1)
        y=random.randrange(ncols-1)
        while world.knock(x,y)!=None:
            x=random.randrange(nrows-1)
            y=random.randrange(ncols-1)
        #print("check")
        child=(Bear(world, (x,y)))
    return world
```

```python
w = wbf(10,10,3,12)
w.grid
```

Output:
```
array([[None, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, Fish, None, None, None],
       [None, None, None, None, Fish, None, None, None, None, None],
       [None, Fish, None, Fish, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [Fish, None, Fish, Bear, None, None, Fish, Bear, None, None],
       [None, None, Fish, None, None, None, None, None, Fish, None],
       [None, Bear, None, None, None, None, None, Fish, None, None],
       [None, None, None, None, None, None, None, None, Fish, None],
       [None, None, None, None, None, None, None, None, None, None]], dtype=object)
```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: wbf_test, points: 3, grade: true, solution: false -->
```python
w = wbf(10,10,3,12)

dct = { Bear: 0, Fish: 0 }
for x in w.biota():
    dct[x.__class__] += 1

assert dct[Bear] == 3
assert dct[Fish] == 12
```

##### <span style="color:red">Experiments</span>

Run some experiments with the top level simulation loop and describe the results in the markdown cell below.  Some ideas of things to try:
* The settings for the Bear class `breed_interval` and `survive_without_food` variables come from the textbook.  Will the world ever run out of bears with these settings?
* Change the settings so `Bear.breed_interval` is larger than `Bear.survive_without_food`.  How does that change the outcome?
* Set the `breed_interval` counter for the Fish class to a smaller number (e.g. 4) so the world has more fish.  What effect does that have?
* Write a function that runs the simulation for a specified number of generations, or until there are no more objects left in the grid.  What combination of parameters leads to the largest number of time steps before the simulation halts.

```python
def step_system(world,lifetime):
    for x in world.biota():
        x.live()
    for x in world.biota():
        x.move()
    for i in range (lifetime):
        world.live()
    return world.grid

step_system(wbf(10,10,3,12),10)

```

Output:
```
array([[None, None, None, None, None, None, None, None, None, None],
       [Fish, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, Fish, None, None],
       [None, None, None, Fish, None, None, None, Fish, None, Fish],
       [Fish, None, None, None, None, None, None, None, None, Bear],
       [None, None, None, None, None, None, None, Fish, None, None],
       [None, None, None, None, None, None, Bear, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, Bear, None],
       [None, None, None, None, None, None, None, None, None, None]], dtype=object)
```

```python
Bear.survive_without_food = 6
#Fish don't starve yet, so I made them live to 100
Fish.survive_without_food=100
step_system(wbf(10,10,4,10),100)
```

Output:
```
array([[None, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [Fish, None, None, None, None, None, None, None, None, None],
       [Fish, None, None, None, Bear, None, None, None, None, None],
       [None, None, None, None, None, None, Fish, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, Fish, None, None, None, None, None],
       [None, None, None, None, None, Fish, Fish, None, None, None],
       [None, None, Fish, None, None, None, Fish, None, None, None]], dtype=object)
```

<!-- grade_id: experiments, points: 10, grade: true, solution: true -->
**************
Due to the breed interval being less then the survive without food amount, the bears will never go extinct in this simulation with the values taken from the book.

If we change the breed interval to be greater then survive without food, the bears die off

When we raise the breed_interval of the fish, they multiply too fast and overcrowd themselves into extinction

As I altered the parameters, this world desplays a certain set of rules for survival. optimal surive without foo is about 6, there must be at least 2 times as many fish for the bears to survive, but much more then that the fish overcrowd. The magic number for simulations was 4 bears and 10 fish, from my exeriments. It always lasted 100 generations, and once 1000. Most other combinations failed to reach 100 generations consitently with both animals still in the biome. 

*************


##  <span style="color:teal;">Projects for Groups</span>

Here are two suggestions for ways to extend the simulation.  We will consider other types of extensions -- send a request to `conery@uoregon.edu` with your proposal.

Groups with three people should implement two extensions.  Groups with two people can choose either extension, or do both for extra credit.

* Implement the Plant class described in Section 10.7 of the textbook, and modify the Fish class so fish eat plants and die if they don't find enough food.

* Implement a GUI using `tkinter` that is similar to the Solar System GUI.  A canvas should display the world along with images for the Bear and Fish objects (you can download `Bear.gif` and `Fish.gif` from Canvas).  Use a spinbox or text entry box to specify the number of Bear and Fish objects and the number of time steps to run.  Include a Run button to start the simulation.


*********
Below is my implementation of the Plant class. I edited my Fish class above to enable the fish to eat plants and die from starvation. I also rewrote the top leval function so i could test the new interactions and test them with timestep. The plants do not die from old age, only from being eaten, so after multile runs (50+) there are only plants left as the fish and bears starve of die off.
************


```python
class Plant(Animal):
    breed_interval=5
    
    def __init__(self, world, tuple):
        Animal.__init__(self, world, tuple)
        
    def live(self):
        self.breedTick+=1
        if (self.breedTick>=Plant.breed_interval):
            self.tryToBreed
    def __repr__(self):
        return "Plnt"
    

```

```python
def wbfp(nrows, ncols, nbears, nfish, nplants):
    '''
    Return a new World object with the specified number of Bear and Fish objects.
    '''
    world=World(nrows,ncols)
    for i in range(nfish):
        
        x=random.randrange(nrows-1)
        y=random.randrange(ncols-1)
        
        while world.knock(x,y)!=None:
            x=random.randrange(nrows-1)
            y=random.randrange(ncols-1)
        child=Fish(world, (x, y))
    for i in range(nbears):
        
        x=random.randrange(nrows-1)
        y=random.randrange(ncols-1)
        while world.knock(x,y)!=None:
            x=random.randrange(nrows-1)
            y=random.randrange(ncols-1)
        #print("check")
        child=(Bear(world, (x,y)))
    for i in range(nplants):
        
        x=random.randrange(nrows-1)
        y=random.randrange(ncols-1)
        while world.knock(x,y)!=None:
            x=random.randrange(nrows-1)
            y=random.randrange(ncols-1)
        #print("check")
        child=(Plant(world, (x,y)))
    return world
```

```python
world=wbfp(10,10,2,2,20)
world.grid
```

Output:
```
array([[None, Plnt, Plnt, None, Fish, None, Plnt, None, None, None],
       [None, Plnt, None, None, Plnt, None, None, None, Plnt, None],
       [None, None, None, None, None, None, None, None, None, None],
       [None, None, None, Plnt, None, None, None, Plnt, None, None],
       [None, None, None, Plnt, Plnt, None, None, None, None, None],
       [Plnt, None, Plnt, Plnt, None, None, Bear, None, Plnt, None],
       [None, Bear, None, None, Plnt, None, None, None, None, None],
       [Plnt, None, None, None, None, None, None, Plnt, None, None],
       [None, None, None, Plnt, Plnt, None, Plnt, None, None, None],
       [None, None, None, None, None, None, None, None, None, None]], dtype=object)
```

```python
step_system(world,10)
```

Output:
```
array([[None, None, None, None, None, None, Plnt, None, Plnt, None],
       [None, None, Plnt, None, None, Plnt, None, None, None, None],
       [None, None, Plnt, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [None, None, Plnt, None, None, None, None, None, None, None],
       [Plnt, None, Plnt, None, None, None, Bear, None, None, None],
       [None, None, Bear, None, None, None, None, Plnt, None, None],
       [Plnt, None, None, Plnt, None, None, None, Plnt, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [Plnt, None, None, Plnt, None, None, Plnt, None, None, None]], dtype=object)
```

```python

```

```python

```