# <span style="color:teal;">CIS 211 Project 5:  Building GUIs with Tk</span>

##### Due 11:00 P.M. Friday Feb 5

Reading:
* [TkDocs - Tk Tutorial](http://www.tkdocs.com/tutorial/index.html)
* [Tkinter Tutorial](http://www.python-course.eu/python_tkinter.php) from `python-course.eu`

### <span style="color:teal">Setup</span>

Execute the following code cell to tell IPython you want to create windows with Tk and to import the `tkinter` library.

```python
%gui tk
import tkinter as tk

```

**Important:** &nbsp; The autograder tests assume `tkinter` has been imported as shown above.  You should follow the same convention, and refer to widget classes as `tk.Label`, `tk.Entry`, and so on.

### <span style="color:teal">1. &nbsp; Hello, World (20 points)</span>

The code cell for this project has an assignment statement that creates a top level Tk application named `hello_app`.  

Add statements to the code cell so the program has a label that says "Hello, World":

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/hello_gui.png"/>

The example above uses 24-point Helvetica.  You can use the default text, but feel free to experiment with other fonts, font sizes, colors, *etc*.

##### <span style="color:red">Documentation:</span>

Use this markdown cell to describe how the "hello" program works.

<!-- grade_id: doc_1, points: 5, grade: true, solution: true -->
************************

We start by importing the Label class for our GUI, then we initializing the hello_app application. Then we add a title and the Label information. Finally, we call pack on our label, and Ipython creates the widget for us.

-Amber Straub

**************************

##### <span style="color:red">Code:</span>

Put your hello program in the following code cell.

<!-- grade_id: hello, points: 5, grade: true, solution: true -->
```python
from tkinter import Tk, Label
hello_app = tk.Tk()
hello_app.title("tk #2")
greeting=Label(hello_app, text = 'Hello, World!')
greeting.pack()
#hello_app.mainloop()

```

##### <span style="color:red">Autograder Tests:</span>

The three tests for this program check to see if the application has a single widget, that the widget is a label, and the label text starts with the string "hello".

**Do not edit, copy, or delete these cells.**

<!-- grade_id: hello_test_1, points: 2, grade: true, solution: false -->
```python
widgets = list(hello_app.children.values())
assert len(widgets) == 1
```

<!-- grade_id: hello_test_2, points: 3, grade: true, solution: false -->
```python
widgets = list(hello_app.children.values())
assert isinstance(widgets[0], tk.Label)
```

<!-- grade_id: hello_test_3, points: 5, grade: true, solution: false -->
```python
widgets = list(hello_app.children.values())
assert widgets[0]['text'].lower().startswith('hello')
```

### <span style="color:teal">2. &nbsp; Payment Calculator (50 points)</span>

Create a program that computes the monthly payment on a loan. The payment is a function of the loan amount, the annual interest rate, and the number of years to pay off the loan. To calculate the payment, first compute two values named $r$ and $p$, defined by these equations:

$ r = ( \textrm{rate} \; / \; 100 ) \; / \; 12 $

$ p = 12 \times \textrm{years} $

The monthly payment is then

$$ \textrm{payment} = \frac{ r \times \textrm{amount} }{ 1 - (1 + r)^{-p} } $$

For example, if you take out a 30-year loan for \$150,000 at an annual rate of 4.5% your payment would be \$760.03 per month.

The GUI for this application should have four text entry boxes, and each box should have a label to its left that explains what the box holds.  The first three boxes are for the user to enter the loan amount, the interest rate, and the number of years to repay the loan. 

The GUI should have a button labeled “compute payment”. When the user clicks this button, your program should read the values of the three parameters, use the equations shown above to calculate the monthly payment, and display the payment amount in the fourth text entry box.

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/payment_gui.png"/>

The code cell for this program has the outline for the program.  First fill in the body of the function named `monthly_payment` and make sure it computes the correct payment value given an amount, rate, and loan period.  When your write the button callback function have it call the `monthly_payment` function.

Next add statements that create the labels, boxes, and button for the application.  When you create the Entry and Button widgets, make sure you save them in global variables with the following names so the autograder tests will succeed:
* `amount`
* `rate`
* `years`
* `payment`
* `button`

##### <span style="color:red">Documentation:</span>

Use this markdown cell to describe your "payment" program.

<!-- grade_id: doc_2, points: 10, grade: true, solution: true -->
*****************************

I started by initiating two functions, one that calculates the monthly payment, and another that extracts values from the user input enteries for insertion into the monthly payment function. Then I initialized each variable we would need to work with, followed by initiallizing and placing the Labels, Enteries and button in the widget. I played with padding and column allignment, to create an astesically pleasing widget.


-Amber Straub

******************

##### <span style="color:red">Code:</span>

Put your payment program in the following code cell.

<!-- grade_id: payment, points: 15, grade: true, solution: true -->
```python
from tkinter import *
def monthly_payment(amt, rate, yrs):
    r=(rate/100)/12
    p=(12*yrs)
    total=((r*amt)/(1-((1+r)**-(p))))
    return total

def compute_payment_cb():
    try:
        payment.delete(0, tk.END)
        amt=float(amount.get())
        Rate=float(rate.get())
        Years=float(years.get())
        final=(format((monthly_payment(amt, Rate, Years)), ".2f"))
        payment.insert(0,"$"+str(final))
    except ValueError:
        payment.insert(0,"Must enter amt, rate, yrs")
        pass

payment_app = tk.Tk()
payment_app.title("tk #3")

amount=StringVar()
rate=StringVar()
years=StringVar()
payment=StringVar()


Label(payment_app, text = 'Loan Amount:', width=10, font = ('Helvectica', 20)).grid(row = 0, column = 0, columnspan=1, padx =5, pady = 50)
Label(payment_app, text = 'Interest Rate:', width=10, font = ('Helvectica', 20)).grid(row = 1, column = 0, columnspan=1, padx = 5, pady = 50)
Label(payment_app, text = 'Loan Period', width=10, font = ('Helvectica', 20)).grid(row = 2, column = 0, columnspan=1, padx =5, pady= 50)
Label(payment_app, text = 'Payment:', width=10, font = ('Helvectica', 20)).grid(row = 3, column = 0, columnspan=1, padx = 5, pady = 50)


amount=tk.Entry(payment_app,width=10, textvariable=amount)
amount.grid(row = 0, column = 1, columnspan=1, padx = 20, pady = 50)

rate=Entry(payment_app,width=10, textvariable=rate)
rate.grid(row = 1, column = 1, columnspan=1, padx = 20, pady = 50)

years=Entry(payment_app,width=10, textvariable=years)
years.grid(row = 2, column = 1, columnspan=1, padx = 20, pady = 50)

payment=Entry(payment_app,width=30, textvariable=payment)
payment.grid(row = 3, column = 1, columnspan=1, padx = 20, pady = 50)

button=tk.Button(payment_app, text="Compute Payment", command=compute_payment_cb)
button.grid(columnspan=2, row=4,padx = 20, pady = 50)



```

##### <span style="color:red">Autograder Tests:</span>

The three tests for this program make sure the `monthly_payment` function computes the correct payment value, check to see if the GUI has the correct number of widgets, and test the GUI by inserting values into the entry boxes and clicking the button.

**Do not edit, copy, or delete these cells.**

<!-- grade_id: payment_test_1, points: 5, grade: true, solution: false -->
```python
assert round(monthly_payment(20000, 3.0, 6), 2) == 303.87
```

<!-- grade_id: payment_test_2, points: 10, grade: true, solution: false -->
```python
counts = { tk.Label: 0, tk.Entry: 0, tk.Button: 0 }

for x in payment_app.children.values():
    counts[type(x)] += 1
    
assert counts == { tk.Label: 4, tk.Entry: 4, tk.Button: 1 }
```

<!-- grade_id: payment_test_3, points: 10, grade: true, solution: false -->
```python
amount.insert(0, '20000')
rate.insert(0, '3.0')
years.insert(0, '6')
button.invoke()

assert payment.get() == '$303.87'
```

### <span style="color:teal">3. &nbsp; Canvas (30 points)</span>

Fill in the application named `canvas_app`.  The application will have just one widget, a Tk canvas with a size of $400 \times 400$ pixels. 

You also need to fill in the body of the function named `make_circle` that will draw a circle on your canvas.  The parameters are a reference to the canvas widget, three integers for the size and location of the circle, and a string with a color name.

Use the function to draw three circles on your canvas.  You can choose whatever size, location, and color you would like.

<img src="http://www.cs.uoregon.edu/Classes/16W/cis211/images/gui/three_circles.png"/>


##### <span style="color:red">Documentation:</span>

Use this markdown cell to describe how the "canvas" program works.

<!-- grade_id: doc_3, points: 5, grade: true, solution: true -->
*************

This program draws three circles in a canvas widget. These circle's grow exponentially bigger, change possition, and change color on each call. 


-Amber Straub

*****************

##### <span style="color:red">Code:</span>

Put your canvas program in the following code cell.

<!-- grade_id: canvas, points: 10, grade: true, solution: true -->
```python
def make_circle(w,z, x, y, color):

    canvas.create_oval((w,z, x, y), fill=color[0])
    
    

canvas_app = tk.Tk()
canvas_app.title("tk #4")
canvas=Canvas(canvas_app)
canvas.grid(column=0, row=0, sticky=(N, W, E, S))

color_list=["red", "green", "blue", "yellow"]

x=20
y=20
w=40
z=40


for i in range (3):
    make_circle(w,z, x, y, color_list)
    x+=120
    y+=120
    w+=60
    z+=60
    color_list=color_list[1:]

```

```python
canvas.create_oval((120,120, 10, 10), fill="red")
```

Output:
```
C:\Users\Amber\Anaconda2\lib\tkinter\__init__.py in create_oval(self, *args, **kw)
   2332     def create_oval(self, *args, **kw):
   2333         """Create oval with coordinates x1,y1,x2,y2."""
-> 2334         return self._create('oval', args, kw)
   2335     def create_polygon(self, *args, **kw):
   2336         """Create polygon with coordinates x1,y1,...,xn,yn."""

C:\Users\Amber\Anaconda2\lib\tkinter\__init__.py in _create(self, itemType, args, kw)
   2317         return self.tk.getint(self.tk.call(
   2318             self._w, 'create', itemType,
-> 2319             *(args + self._options(cnf, kw))))
   2320     def create_arc(self, *args, **kw):
   2321         """Create arc shaped region with coordinates x1,y1,x2,y2."""

TclError: invalid command name ".79781136"
```

##### <span style="color:red">Autograder Tests:</span>

The tests for this program make sure the application has a single canvas widget and that the canvas has three circles.

**Do not edit, copy, or delete these cells.**

<!-- grade_id: canvas_test_1, points: 5, grade: true, solution: false -->
```python
widgets = list(canvas_app.children.values())
assert len(widgets) == 1 and isinstance(widgets[0], tk.Canvas)
```

<!-- grade_id: canvas_test_2, points: 10, grade: true, solution: false -->
```python
widgets = list(canvas_app.children.values())
for n in widgets[0].find_all():
    assert widgets[0].type(n) == 'oval'
```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```